rule_files:
  - crabcache.rules

evaluation_interval: 1m

tests:
- interval: 1m
  input_series:
  - series: 'crabcache_exporter_num_cpus{env="prod",instance="test-instance",host="k8s-test"}'
    values: '0+0x100000'
  alert_rule_test:
      - eval_time: 10m
        alertname: Crabcache is down
        exp_alerts:
            - exp_labels:
                 severity: high
                 tag: cmsweb
                 service: crabcache
                 host: k8s-test
                 action: Please restart Crabcache on test-instance
                 instance: test-instance
                 env: prod
              exp_annotations:
                 summary: "crabcache prod is down"
                 description: "prod has been down for more than 5m"

- interval: 1m
  input_series:
  - series: 'avg_mem_pct{env="prod",host="k8s-test"}'
    values: '100+100x7'
  alert_rule_test:
      - eval_time: 10m
        alertname: crabcache high memory usage
        exp_alerts:
            - exp_labels:
                 severity: high
                 tag: cmsweb
                 service: crabcache
                 host: k8s-test
                 action: Please restart crabcache service
                 env: prod
              exp_annotations:
                 summary: "crabcache prod has high memory usage"
                 description: "prod memory usage above 90% for more than 5m"

- interval: 1m
  input_series:
  - series: 'avg_cpu_pct{env="prod",host="k8s-test"}'
    values: '100+100x7'
  alert_rule_test:
      - eval_time: 10m
        alertname: crabcache high cpu usage
        exp_alerts:
            - exp_labels:
                 severity: medium
                 tag: cmsweb
                 service: crabcache
                 host: k8s-test
                 env: prod
                 action: Please restart crabcache service
              exp_annotations:
                 summary: "crabcache prod has high CPU usage"
                 description: "prod CPU usage above 90% for more than 5m"
