rule_files:
  - crabserver.rules

evaluation_interval: 1m

tests:
- interval: 1m
  input_series:
  - series: 'crabserver_num_cpus{env="prod",instance="test-instance",host="k8s-test"}'
    values: '0+0x100000'
  alert_rule_test:
      - eval_time: 10m
        alertname: CRAB server is down
        exp_alerts:
            - exp_labels:
                 severity: high
                 tag: cmsweb
                 service: crab
                 host: k8s-test
                 action: Please restart CRAB server on test-instance
                 instance: test-instance
                 env: prod
              exp_annotations:
                 summary: "crabserver prod is down"
                 description: "prod has been down for more than 5m"
- interval: 1m
  input_series:
  - series: 'avg_open_fds_8m{env="prod",instance="test-instance",host="k8s-test"}'
    values: '51+1x100'
  alert_rule_test:
      - eval_time: 10m
        alertname: CRAB server service has high number of fds
        exp_alerts:
            - exp_labels:
                 severity: high
                 tag: cmsweb
                 service: crab
                 host: k8s-test
                 action: Please restart CRAB server on test-instance
                 instance: test-instance
                 env: prod
              exp_annotations:
                 summary: "CRAB prod environment"
                 description: "prod has high level of fds 61 (avg 8m) for more than 1m"
