rule_files:
  - reqmgr2.rules

evaluation_interval: 1m

tests:
- interval: 1m
  input_series:
  - series: 'reqmgr2_num_cpus{env="prod",host="k8s-test",instance="test-instance"}'
    values: '0+0x100'
  alert_rule_test:
      - eval_time: 10m
        alertname: ReqMgr2 service is down
        exp_alerts:
            - exp_labels:
                 severity: high
                 tag: cmsweb
                 service: reqmgr2
                 host: k8s-test
                 kind: dmwm
                 action: Please restart ReqMgr2 service on test-instance
                 instance: test-instance
                 env: prod
              exp_annotations:
                 summary: "reqmgr2 prod is down"
                 description: "prod has been down for more than 5m"

- interval: 1m
  input_series:
  - series: 'reqmgr2_proc_mem{env="prod",host="k8s-test"}'
    values: '100+100x7'
  alert_rule_test:
      - eval_time: 10m
        alertname: ReqMgr2 high memory usage
        exp_alerts:
            - exp_labels:
                 severity: medium
                 tag: cmsweb
                 service: reqmgr2
                 host: k8s-test
                 kind: dmwm
                 env: prod
              exp_annotations:
                 summary: "reqmgr2 prod has large memory usage"
                 description: "prod memory usage above 70% for more than 5m"

- interval: 1m
  input_series:
  - series: 'reqmgr2_proc_cpu{env="prod",host="k8s-test"}'
    values: '100+100x7'
  alert_rule_test:
      - eval_time: 10m
        alertname: ReqMgr2 high cpu usage
        exp_alerts:
            - exp_labels:
                 severity: medium
                 tag: cmsweb
                 service: reqmgr2
                 host: k8s-test
                 kind: dmwm
                 env: prod
              exp_annotations:
                 summary: "reqmgr2 prod has large cpu usage"
                 description: "prod cpu usage above 70% for more than 5m"
